# Principles of Spatial Date Mining and Machine WS2022/2023 
# Lectuer team: Martin Werner, Hao Li

## Regression Excercise Part 1

In this exercise, we will try to implement linear regression and polynomial regression with the Boston housing price dataset, which we have used as an example during the lecutre. 

First, we would load the datasets together, do some basic data exploration, and split the entire datasets into train and test. 

Then, there are mainly two tasks for you in this excercise:

- Build the linear regression model between 'RM' (average Number of rooms per person) and 'MEDV' (housing prices in 1000s US dollars), and calculate the mean_squared_error in the test dataset.
    
- Build the polynomial regression model between 'LSTAT' (percentage of lower status of the population) and 'MEDV' using different degrees of polynominal function (i.e., x^2, x^3 etc), and calculate the mean_squared_error in the test dataset.

## Example Code:

```python
!pip install scikit-learn==1.1.1 -i https://pypi.tuna.tsinghua.edu.cn/simple
```

Output:
```
Defaulting to user installation because normal site-packages is not writeable
Looking in indexes: https://pypi.tuna.tsinghua.edu.cn/simple
Collecting scikit-learn==1.1.1
  Downloading https://pypi.tuna.tsinghua.edu.cn/packages/43/bc/7130ffd49a1cf72659c61eb94d8f037bc5502c94866f407c0219d929e758/scikit_learn-1.1.1-cp310-cp310-manylinux_2_17_x86_64.manylinux2014_x86_64.whl (30.4 MB)
[2K     [90m━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━[0m [32m30.4/30.4 MB[0m [31m1.7 MB/s[0m eta [36m0:00:00[0m00:01[0m00:01[0mm
[?25hRequirement already satisfied: scipy>=1.3.2 in /home/instructor/.local/lib/python3.10/site-packages (from scikit-learn==1.1.1) (1.10.0)
Requirement already satisfied: numpy>=1.17.3 in /home/instructor/.local/lib/python3.10/site-packages (from scikit-learn==1.1.1) (1.24.2)
Requirement already satisfied: joblib>=1.0.0 in /home/instructor/.local/lib/python3.10/site-packages (from scikit-learn==1.1.1) (1.2.0)
Requirement already satisfied: threadpoolctl>=2.0.0 in /home/ins
```

```python
import numpy as np
import pandas as pd
import operator

#imports from sklearn library
from sklearn import datasets
from sklearn.linear_model import LinearRegression
from sklearn.model_selection import train_test_split, cross_val_score
from sklearn.metrics import mean_squared_error
from sklearn.preprocessing import PolynomialFeatures
from scipy.optimize import curve_fit

#Visualization Libraries
import matplotlib.pyplot as plt


#To plot the graph embedded in the notebook
%matplotlib inline
```

### Load the dataset

```python
# loading the Bouston dataset direclty from sklearn
boston = datasets.load_boston()
```

Output:
```
Cell In[31], line 2
      1 # loading the Bouston dataset direclty from sklearn
----> 2 boston = datasets.load_boston()

File ~/.local/lib/python3.10/site-packages/sklearn/datasets/__init__.py:156, in __getattr__(name)
    105 if name == "load_boston":
    106     msg = textwrap.dedent(
    107         """
    108         `load_boston` has been removed from scikit-learn since version 1.2.
   (...)
    154         """
    155     )
--> 156     raise ImportError(msg)
    157 try:
    158     return globals()[name]

ImportError: 
`load_boston` has been removed from scikit-learn since version 1.2.

The Boston housing prices dataset has an ethical problem: as
investigated in [1], the authors of this dataset engineered a
non-invertible variable "B" assuming that racial self-segregation had a
positive impact on house prices [2]. Furthermore the goal of the
research that led to the creation of this dataset was to study the
impact of air quality but it did not give adequate demonstration of the
```

### Now let us print the attributes,, and check which attributes are available

```python
# Print the python data tpye and size of boston dataset together with the feature_names

print(type(boston))
print('\n')
print(boston.keys())
print('\n')
print(boston.data.shape)
print('\n')
print(boston.feature_names)

```

Output:
```
<class 'str'>


AttributeError                            Traceback (most recent call last)
Cell In[29], line 5
      3 print(type(boston))
      4 print('\n')
----> 5 print(boston.keys())
      6 print('\n')
      7 print(boston.data.shape)

AttributeError: 'str' object has no attribute 'keys'
```

## The features can be summarized as follows (https://www.cs.toronto.edu/~delve/data/boston/bostonDetail.html):



- CRIM: This is the per capita crime rate by town
- ZN: This is the proportion of residential land zoned for lots larger than 25,000 sq.ft.
- INDUS: This is the proportion of non-retail business acres per town.
- CHAS: This is the Charles River dummy variable (this is equal to 1 if tract bounds river; 0 otherwise)
- NOX: This is the nitric oxides concentration (parts per 10 million)
- RM: This is the average number of rooms per dwelling
- AGE: This is the proportion of owner-occupied units built prior to 1940
- DIS: This is the weighted distances to five Boston employment centers
- RAD: This is the index of accessibility to radial highways
- TAX: This is the full-value property-tax rate per 10,000 dollars
- PTRATIO: This is the pupil-teacher ratio by town
- B: This is calculated as 1000(Bk — 0.63)², where Bk is the proportion of people of African American descent by town
- LSTAT: This is the percentage lower status of the population
- MEDV: This is the median value of owner-occupied homes in 1,000 dollars



```python
# Convert boston data into pandas dataframe 

boston_data = pd.DataFrame(boston.data, columns=boston.feature_names)
boston_data.head()
```

Output:
```
AttributeError                            Traceback (most recent call last)
Cell In[30], line 3
      1 # Convert boston data into pandas dataframe 
----> 3 boston_data = pd.DataFrame(boston.data, columns=boston.feature_names)
      4 boston_data.head()

AttributeError: 'str' object has no attribute 'data'
```

```python
# Append the price feature as a colunm from the boston
boston_data['MEDV'] = boston.target
print(boston_data)
```

```python
# scatter plot using'RM' and 'LSTAT' as x axis, and 'MEDV' as y axis

plt.figure(figsize=(16, 8))

features = ['LSTAT', 'RM']
target = boston_data['MEDV']

for i, col in enumerate(features):
    plt.subplot(1, len(features) , i+1)
    x = boston_data[col]
    y = target
    plt.scatter(x, y, marker='x')
    plt.title(col)
    plt.xlabel(col)
    plt.ylabel('MEDV')
```

# Now, it is your turn!

## Task No.1: fit a linear regression model 

### Q1.1: define the split ratio between training and validation data

<!-- grade_id: cell-36bf81b3cc459deb, points: 0, grade: false, solution: true -->
```python
# split the boston hoursing dataset into training and validation

# tips: try to remember the default ratio we disucssed in the lecture 
ratio = 0.2 # split ratio

X = boston_data['RM'].values.reshape(-1,1)

Y = boston_data['MEDV']

### BEGIN SOLUTION
x_train, x_test, y_train, y_test = train_test_split(X, Y, test_size = ratio, random_state=5)

#print the size of train and train dataset
print(x_train.shape)
print(y_train.shape)
### END SOLUTION
```

<!-- grade_id: cell-68aa8fe9670b6419, points: 2, grade: true, solution: false -->
```python
try: 
   assert x_train.shape == (404, 1)
except: 
   print('Element shape of x_train is not correct. shape is not correct')

try: 
   assert y_train.shape == (404,)
except: 
   print('Element shape of y_train is not correct. shape is not correct')

```

### Q1.2: print the size of test dataset

```python
### BEGIN SOLUTION
print(x_test.shape)
print(y_test.shape)
### END SOLUTION
```

<!-- grade_id: cell-6386ce02461d45e3, points: 1, grade: true, solution: false -->
```python
try: 
   assert x_test.shape == (102, 1)
except: 
   print('Element shape of x_test is not correct. shape is not correct')

try: 
   assert y_test.shape == (102,)
except: 
   print('Element shape of y_test is not correct. shape is not correct')

```

```python
# initial a linear regression model
lin_reg = LinearRegression()

# fit a linear regression to linear features.
# tips: for the LinearRegression.fit(x,y) function, herein x will be your input feature and y is the label
lin_reg.fit(x_train, y_train)

# calculate the y_hat from the trained model
y_train_pred = lin_reg.predict(x_train)
```

### Q1.3:  plot the fitted line

<!-- grade_id: cell-775575b027be2af9, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
# plot the results
plt.figure(figsize=(8, 8))
plt.scatter(x_train, y_train, marker='x')

# sort the values of x before line plot
sort_axis = operator.itemgetter(0)
sorted_zip = sorted(zip(x_train,y_train_pred), key=sort_axis)
x_plot, y_polt_pred = zip(*sorted_zip)

plt.plot(x_plot, y_polt_pred, color='red')
### END SOLUTION
```

<!-- grade_id: cell-e59ed7ab8d73f8eb, points: 1, grade: true, solution: false -->
```python
try: 
   assert len(x_train) == len(sorted_zip) == 404

except:  
   print("Incorrcet y_pred. Please check your prediction result again.")

```

```python
# tips: calculate the y_hat for the validatiaon set in a similar way.
y_test_predict = lin_reg.predict(x_test)

# model evaluation for validatiaon set
mse = mean_squared_error(y_test, y_test_predict)
print("The model performance for validatiaon set")
print("--------------------------------------")
print('MSE is {}'.format(mse))
```

## Task No.2: fit a polynomial regression model

```python
# split the boston hoursing dataset into training and testing 

ratio = 0.2 # split ratio

X = boston_data['LSTAT'].values.reshape(-1,1)

Y = boston_data['MEDV']


```

### Q2.1: split the training data and the test data

<!-- grade_id: cell-0a18567238b5e95f, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
# tips: similar as we did for linear feature
x_train, x_test, y_train, y_test = train_test_split(X, Y, test_size = 0.2, random_state=5)
### END SOLUTION
```

<!-- grade_id: cell-4be4a6bead75bc24, points: 1, grade: true, solution: false -->
```python
try: 
   assert x_train.shape ==(404, 1) and y_train.shape ==(404,) and x_test.shape == (102, 1) and y_test.shape == (102,)
except: 
   print("The result seems to be different. Try again with correct parameters.") 
```

```python
#print the size of train and test dataset
print(x_train.shape)
print(y_train.shape)
print(x_test.shape)
print(y_test.shape)
```

```python
# this the degree of polynomial feature
poly_degree = 2

# build the polynomial feature space
poly = PolynomialFeatures(degree=poly_degree, include_bias=False)
x_train_poly = poly.fit_transform(x_train)

# initial a linear regression model
lin_reg = LinearRegression()

# fit a linear regression to polynomial features.
# tips: now you want to fit the model with the polynomical features you get from the previous step
lin_reg.fit(x_train_poly, y_train)
y_train_poly_pred = lin_reg.predict(x_train_poly)
```

```python
# plot the results
plt.figure(figsize=(8, 8))
plt.scatter(x_train, y_train, marker='x')
# sort the values of x before line plot
sort_axis = operator.itemgetter(0)
sorted_zip = sorted(zip(x_train,y_train_poly_pred), key=sort_axis)
x_plot, y_poly_pred = zip(*sorted_zip)
# plot the fitted line
plt.plot(x_plot, y_poly_pred, color='red')
```

### Q2.2:  print the result of residulas (MSE)

<!-- grade_id: cell-9c0801fb4e12ac2a, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
# model evaluation for testing set
# tips: now do the same polynomial transformation for validation set and calculat the MSE
x_test_poly = poly.fit_transform(x_test)
y_test_predict = lin_reg.predict(x_test_poly)
mse = mean_squared_error(y_test, y_test_predict)
### END SOLUTION
```

```python
print("The model performance for validatiaon set")
print("--------------------------------------")
print(mse)
```

<!-- grade_id: cell-d01a6de436464940, points: 1, grade: true, solution: false -->
```python
try: 
   assert  20 <= mse <= 29
except: 
   print("The result seems to be different. Try again with correct parameters.") 

```

# Well Done! Now you know how to train a linear regression model!

## Addition task No.3:: How about a exponential feature space?

```python
# split the boston hoursing dataset into training and testing 

ratio = 0.2 # split ratio

X = boston_data['LSTAT']

Y = boston_data['MEDV']

x_train, x_test, y_train, y_test = train_test_split(X, Y, test_size = 0.2, random_state=5)

# print the size of train and test dataset

print(x_train.shape)
print(y_train.shape)
print(x_test.shape)
print(y_test.shape)
```

```python
def func_exp(x, a, b, c):
        #c = 0
        return a * np.exp(b * x) + c

def exponential_regression(x_data, y_data):
    popt, pcov = curve_fit(func_exp, x_data, y_data, p0 = (-1, 0.01, 0))
    print(popt)
    plt.figure(figsize=(8, 8))
    puntos = plt.plot(x_data, y_data, 'x', label = "data")
    # tips: create exponential feature sapce using the aforedefined two functions: def func_exp and def exponential_regression
    y_exp_pred = func_exp(x_data, *popt)
    # sort the values of x before line plot
    sort_axis = operator.itemgetter(0)
    sorted_zip = sorted(zip(x_data,y_exp_pred), key=sort_axis)
    x_data, y_exp_pred = zip(*sorted_zip)
    curva_regresion = plt.plot(x_data, y_exp_pred , color='red')
    plt.legend()
    plt.show()
```

### Q3.1: apply the regression function to the data 

<!-- grade_id: cell-4ad51d1bbe68ee06, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
# tips: then fit a exponential regression model
exponential_regression(x_train, y_train)
### END SOLUTION
```

<!-- grade_id: cell-60c060bf8c5bda69, points: 1, grade: true, solution: false -->
```python
try: 
   assert x_train.shape == y_train.shape == (404,) and x_test.shape == y_test.shape == (102,)
except: 
   print("Incorrect usage of exponential_regression function")

```